Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Problem 1 - Testing for Fibonacci Number (10 points)

Develop functions to handle [Fibonacci numbers](http://mathworld.wolfram.com/FibonacciNumber.html). Try to design the quickest possible algorithm and record the time taken for the code to run.

Test if a number belongs to the Fibonacci series, using `if-else` conditions. The logic of testing if a number, $n$, belongs to the Fibonacci series is: if one or both of $(5*n^2 + 4)$ or $(5*n^2 – 4)$ is a perfect square.

Here is what you need to do:
1. You will need to implement the function `fibonacciGenerator`, which takes no inputs. The function should return a Python `list` of length 25, with first 25 Fibonacci numbers.
2. You will need to define a function `fibonacciTester`, which takes a single number as input, and returns `true` if the number belongs to the Fibonacci series, and `false` otherwise. 

<!-- grade_id: cell-42dfa70d19792ff1, points: 0, grade: false, solution: true -->
```python
def fibonacciGenerator():
    """
    Takes no inputs. Retun a list of fist 25 Fibonacci numbers.
    """
    seq = [0,1]
    for i in range (1,24):
        a = seq[i] + seq[i-1]
        seq.append(a)
    ###print ("Fibonacci sequence" , seq)
    ####print (len(seq))
    return seq;
    raise NotImplementedError()
    
# import time
# st = time.time()
# l = fibonacciGenerator()
# en = time.time()
# print(l)
# print("The code took %d seconds."%(en - st))  
```

<!-- grade_id: cell-0ee383f8bfb924f6, points: 5, grade: true, solution: false -->
```python
assert fibonacciGenerator() == [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181, 6765, 10946, 17711, 28657, 46368]

```

<!-- grade_id: cell-4d4607e810cebe9b, points: 0, grade: false, solution: true -->
```python
def fibonacciTester(number):
    """
    The function has one input parameter number, return True or False
    """
    n=((5*(number**2)+4)**.5).is_integer() or ((5*number**2-4)**.5).is_integer()
    if number <0:
        return False
    else:
        return ((5*(number**2)+4)**.5).is_integer() or ((5*number**2-4)**.5).is_integer()


    raise NotImplementedError()

# import time
# st = time.time()
# for i in range(1,21):
#     v = fibonacciTester(i)
#     print(v)
# en = time.time()
# print("The code took %d seconds."%(en - st))    
```

<!-- grade_id: cell-86b24d18aaf6c0b2, points: 5, grade: true, solution: false -->
```python
assert fibonacciTester(0) == True
assert fibonacciTester(4) == False
assert fibonacciTester(-1) == False
assert fibonacciTester(7.1) == False
```

# Problem 2 - Handling Exceptions (10 Points)
In this problem you will read from a text file (**data_p2.txt**) included in the assignment. Please download it and store it in the same directory as this notebook. The code for reading data from the file is given below. You need to store the data and perform a simple mathematical operation on each data point. Each line is _supposed_ to contain a floating point number. But what you will observe is that some lines might have erroneous entries. You need to ignore those lines (Hint: Use `Exception` handling). 

The idea is to implement a function, `computeMedian()`, which reads in a file and computes the median of the numbers and returns the output. You may use the inbuilt function `sort` when computing the median.

_DO NOT USE ANY INBUILT OR OTHER FUNCTION TO DIRECTLY COMPUTE MEDIAN_

<!-- grade_id: cell-a419fd9556586435, points: 0, grade: false, solution: true -->
```python
# Reading from a file. The following code just reads in every line in a text file without doing anything with it.
# You can use the code in your solution. DO NOT CALL this function in your implementation.
def dummyFileReader():
    
    open_file  = open('data_p2.txt').read().splitlines()
    
    
    
def computeMedian():
    
    new_file = []
    
    open_file  = open('data_p2.txt').read().splitlines()
    
    for items in open_file:
        items = items.strip()
        try:
            num = float(items)
            new_file.append(num)
            new_file = sorted(new_file)
        except:
            pass
    median = int((len(new_file)/2))
    #print(median)
    return new_file[median]

```

<!-- grade_id: cell-a5072a0203f9b9eb, points: 10, grade: true, solution: false -->
```python
assert computeMedian() == 0.499675
```

# Problem 3 - Benchmarking Dictionary Performance (10 Points)
We want to check the peformance of the Python inbuilt `del` routine for dropping elements from a dictionary. The question that we want answer for is: 
_Does the time required to delete an element from a Python dictionary depend on the size of the dictionary and on the location of the target element?_

For this experiment, we want to write a function `monitorTime()`, with one parameter input, $n$, that can:

    In each step of the experiment do the following: 
        - Create a dictionary of integers from 1 to N, with your own desired keys. 
        - Within a loop, delete the $i^{th}$ element. Measure the time taken to delete the element (Hint: use the `time` package in Python).
    
    And return the measured time in a list.

Conduct the above experiment for $N = 1,000,000$ and $i$ ranging from $1$ to $1,000,000$ in steps of $10,000$. Make sure that you are recreating the Python dictionary every time. Save the times in a list (let us name it `timeMillion`).

Repeat the above experiment, but set $N = 10,000,000$. This time, vary $i$ from $1$ to $10,000,000$ in steps of $100,000$. Save the times for this experiment in another list, say, `time10Million`.

Plot the two lists using the `matplotlib.pyplot` package. See below for a sample code. You can plot both lists in the same chart to compare.

Using the plots, answer the following (create a markup cell after your code cell(s) and enter your responses):
1. Does the cost of deleting an element depend on the dictionary size?
2. Does the cost of deleting an element depend on the location of the deleted element?

```python
# This cell is only for demonstration purposes, you can reuse parts of this code to create your own plots
# Using matplotlib for plotting
import matplotlib.pyplot as plt
%matplotlib inline 
# above line is needed to make sure that the plots are displayed in the notebook.

#generate some dummy data
x = list(range(0, 20)) # generate a list
y = [x1**2+2*x1 for x1 in x] # create a new list using the original list. Note the Pythonic inline creation of lists.
y1 = [x1**2.3 for x1 in x]
plt.plot(x,y)
plt.plot(x,y1)
```

<!-- grade_id: cell-2220c7af0b65826d, points: 0, grade: false, solution: true -->
```python

def monitorTime(N):
    """
    The function has one input parameter N, N will either be 1,000,000 or 10,000,000.
    """

    import time
    N1 = 1000000
    N2 = 10000000
    i1 = 10000
    i2 = 100000
    step1 = 1 #
    step2 = 1
    i1_location = []
    i2_location = []
    timeMillion=[]
    time10Million=[]
    d1=dict()
    d2=dict()
    if N ==N1:
        
        while step1 <= 100:# 100
            timeslice=[]
            interation = 1 #up to 20
            while interation <=20: #20
                for x in range(1,N1+1):
                    d1[x]=2*x
                start_time = time.time()
                #return (len(d1/))
                del(d1[i1*step1])
                timeslice.append(time.time() - start_time)
                #print(timeslice)
            #print("inte:",interation)
                interation = interation + 1   
            i1_location.append(i1*step1)
            timeMillion.append(sum(timeslice)/20)
        #print("step",step1)
            step1 = step1 + 1 
            
        return numpy.array(timeMillion)
    elif N==N2:
        while step2 <= 100:# 100
            timeslice=[]
            interation = 1 #up to 20
            while interation <=20: #20
                for x in range(1,N2+1):
                    d2[x]=2*x
                start_time = time.time()
                del(d2[i2*step2])
                timeslice.append(time.time() - start_time)
                #print("inte:",interation)
                interation = interation + 1       
            i2_location.append(i2*step2)
            time10Million.append(sum(timeslice)/20)
            #print("step",step2)
            step2 = step2 + 1 
        
        return numpy.array(time10Million)
```

<!-- grade_id: cell-dc4b7724b0cb8035, points: 5, grade: true, solution: false -->
```python
assert len(monitorTime(1000000)) == 100
assert len(monitorTime(10000000)) == 100
```

```python
#Q1
print("The time difference of finding 100 elements from these two scales of datasets are very insignificant that can be ingnored:",sum(numpy.array(timeMillion)-numpy.array(time10Million)))
print("This result represents the set of comparison time that has the biggest difference which is also extremly small:",max(numpy.array(timeMillion)-numpy.array(time10Million)))

```

```python
#Q2
print("The difference between the max and min time in timeMillion list:", max(monitorTime(1000000)-min(monitorTime(10000000)))
print("The difference between the max and min time in time10Million list:", max(monitorTime(1000000))-min(monitorTime(10000000))
```

### 1. Plot your monitored time lists

```python
x = list(range(1, 101))
plt.plot(x,monitorTime(1000000))
plt.plot(x,monitorTime(10000000))
```

<!-- grade_id: cell-b2fa6514d26fe50a, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

```python
# No It does not matter the time or the 
```

# Problem 4 - The Two Envelopes Problem (20 Points)
The Two Envelopes problem is an interesting decision theory problem.

> Consider that you are given two indistinguishable envelopes, each containing money, one contains twice as much as the other. You may pick one envelope and keep the money it contains. Having chosen an envelope at will, but before inspecting it, you are given the chance to switch envelopes. Should you switch?

> The game is this: "stick" or "switch"; It seems obvious that there is no point in switching envelopes as the situation is symmetric. However, because you stand to gain twice as much money if you switch while risking only a loss of half of what you currently have, it is possible to argue that it is more beneficial to switch. The problem is to show what is wrong with this argument.

**Two Envelopes problem**: Implement a function, called `simulateProblem()`, that does the game simulation for the two envelopes problem. Run the simulation 1000 times to figure out the empirical (observed) probability of gaining more money when switching and gaining more money when sticking to the original choice. Each simulation operates as follows:
1. First, randomly pick an envelopes configuration out of the two possible configurations, $(A,2A)$ or $(2A,A)$. In the first configuration, the second envelope has twice the money and in the second configuration, the first envelope has twice the money.
2. Next, randomly pick one of the two envelopes. 
3. Finally, randomly choose to either stick or switch. The program checks if you won (the envelope that picked has more money) or not (the envelope that picked has less money). In case of winning, record if the winning was because of _sticking_ or _switching_.

You can perform the _random_ choice as follows, using the `np.random.randint()` method.
```python
import numpy as np
print(np.random.randint(2))
```
The `simulateProblem()` function takes no arguments and returns two values, first is a boolean output which is `True` if you win and `False` if you lose. In case of a win, the second output is `True` if the win was due to _sticking_ or the lose was due to _switching_ and `False` if the win was due to _switching_ or the lose was due to _sticking_.

Once the method `simulateProblem()` that does the above steps and returns "sticking",or "switching", depending on the win/loss scenario, run the method 1000 times and count the number of times the win was due to _sticking_ to the pick in Step 2, and number of times the win was due to _switching_ the envelope.

<!-- grade_id: cell-170a92182f1bd556, points: 0, grade: false, solution: true -->
```python
import numpy as np

def simulateProblem():
    # YOUR CODE HERE
    configuration = np.random.randint(2)+1
    # 1 stands for A,2A and 2 stands for 2A, A
    envelopePicked = np.random.randint(2)+1
    stickOrSwitch = np.random.randint(2)+1
    # 1 = stick and 2 = switch 
        # Case 1 - When envelope Picked is 2 and configuration is A,2A
    if envelopePicked == 2 and configuration == 1:
         if stickOrSwitch == 1:
            ans = [True, True]
            return ans
         else:
            ans = [False, False]
            return ans
        
                
    # Case 2 - When envelopePicked is 2 and configration is 2A,A 
    if envelopePicked == 2 and configuration == 2:
        if stickOrSwitch == 2:
            ans = [True, False]
            return ans
            #return True, False
        else:
            ans = [False, False]
            return ans
            #return False, False
                
     # Case 3 - When envelopePicked is 1 and configration is A,2A 
    if envelopePicked == 1 and configuration == 1:
        if stickOrSwitch == 2:
            ans = [True, False]
            return ans
            #return True, False
        else:
            ans = [False, False]
            return ans
            #return False,False
                
    # Case 4 - When envelopePicked is 1 and configration is 2A,A 
    if envelopePicked == 1 and configuration == 2:
        if stickOrSwitch == 1:
                ans = [True, True]
                return ans
                #return True, True
        else:
                ans = [False, False]
                return ans

```

```python

```

<!-- grade_id: cell-d307ddac8ccf9503, points: 8, grade: true, solution: false -->
```python
assert type(simulateProblem()[0]) == bool
assert type(simulateProblem()[1]) == bool
```

<!-- grade_id: cell-d540e69cd7d07c9b, points: 0, grade: false, solution: true -->
```python
def run_simulation():
    """
    The function Run the simulation 1000 times to figure out 
    the empirical (observed) probability of gaining more money when switching 
    and gaining more money when sticking to the original choice.
    Return the probability of win due to sticking and win due to switching
    """
    stickCount = 0
    switchCount = 0
    for i in range(1000):
        list = simulateProblem()
        if (list[0] == True):
            if (list[1] == True):
                stickCount += 1
            else:
                switchCount += 1
            print(stickCount,switchCount)
```

### 1. Print you returned values here, your return values should be decimal values

```python
run_simulation()
```

Output:
```
0 1
0 2
0 3
1 3
1 4
2 4
2 5
3 5
3 6
4 6
5 6
6 6
6 7
6 8
6 9
7 9
8 9
8 10
9 10
9 11
9 12
9 13
10 13
11 13
12 13
12 14
12 15
12 16
12 17
13 17
13 18
14 18
15 18
15 19
16 19
17 19
18 19
18 20
18 21
18 22
18 23
19 23
19 24
19 25
20 25
21 25
21 26
22 26
22 27
22 28
23 28
24 28
24 29
25 29
25 30
26 30
26 31
26 32
27 32
28 32
28 33
28 34
28 35
29 35
29 36
30 36
31 36
32 36
33 36
34 36
35 36
35 37
35 38
36 38
37 38
37 39
37 40
38 40
38 41
38 42
38 43
38 44
38 45
38 46
38 47
39 47
40 47
40 48
41 48
41 49
42 49
43 49
44 49
44 50
44 51
45 51
46 51
46 52
46 53
47 53
48 53
48 54
49 54
49 55
50 55
51 55
52 55
52 56
53 56
54 56
55 56
56 56
57 56
57 57
57 58
58 58
59 58
59 59
59 60
59 61
59 62
60 62
61 62
61 63
61 64
61 65
62 65
63 65
64 65
64 66
65 66
65 67
65 68
66 68
66 69
67 69
68 69
68 70
69 70
70 70
70 71
71 71
72 71
73 71
74 71
75 71
75 72
76 72
77 72
78 72
79 72
79 73
80 73
80 74
81 74
82 74
83 74
83 75
83 76
84 76
84 77
85 77
85 78
85 79
85 80
86 80
87 80
88 80
88 81
89 81
90 81
91 81
91 82
9
```

### 2. Give your observation and argument:
 Which is better - Sticking or Switching ？

```python

# there are more wins by sticcking (numbers on the left) vs switching(numbers on the right)
```